<a href="https://colab.research.google.com/github/gsgill7/gsgill7/blob/main/Exercises/06%20Exercises.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Exercise 06.1 (selecting and passing data structures)

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high.

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc, to develop a new version of the function with an improved interface (the interface is the arguments that are passed to the function) with fewer arguments. Add appropriate checks inside your function on the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
import math

x0, y0 = 0.0, 0.0
x1, y1 = 0.0, 2.0
x2, y2 = 3.0, 0.0
p0=(x0,y0)
p1=(x1,y1)
p2=(x2,y2)
coords = (p0,p1,p2)
print(coords[2][0])
def area(coords):
    return abs((coords[0][0]*(coords[1][1]-coords[2][1]) + coords[1][0]*(coords[2][1]-coords[0][1]) + coords[2][0]*(coords[0][1]-coords[1][1]))/ 2)

A = area(coords)
print(A)
assert math.isclose(A, 3.0)
```

Output:
```
3.0
3.0
```

## Exercise 06.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that, given only the vertex coordinates, computes the area of a simple polygon. Test your function for some simple shapes. Pay careful attention to the range of any loops. **Test for polygons without a vertex at the origin.**

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python
square = [(0,0),(1,0),(1,1),(0,1)]
triangle = [(0,0),(1,0),(1,1)]
def area(points):
  area=0
  for n, point in enumerate(points):
    area += 0.5*((points[n][0]*points[(n+1)%len(points)][1])-(points[(n+1)%len(points)][0]*points[n][1]))
  #print(area)
  return(area)
print(area(square))
print(area(triangle))
```

Output:
```
1.0
1.0
0.5
0.5
```

## Exercise 06.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match. Print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(x, y):
    """Return sum of two vectors"""
    if len(x) != len(y):
      print("size mismatch")
      return
    z = [0]*len(x)
    for i in range(len(x)):
      z[i] = x[i] + y[i]
    return z
```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
# tests ##
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Optional (advanced)

Try writing a one-line version that performs the `sum_vector` operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
...
```

## Exercise 06.4 (dictionaries)

Create a dictionary that maps college names (the key) to college abbreviations for at least 5 colleges
(you can find abbreviations at https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges).
From the dictionary, produce and print

1. A Python dictionary from college abbreviation to name; and
1. A Python list of college abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps college names (the key) to dictionaries of:

- College abbreviation
- Year of foundation
- Total number students

for at least 5 colleges. Take the data from https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges. Using this dictionary,

1. Find the college with the greatest number of students and print the abbreviation; and
2. Find the oldest college, and print the number of students and the abbreviation for this college.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
college_abbr = {"Trinity Hall" : "TH", "Gonville and Caius":"CAI", "Corpus Christi" : "CC", "Clare": "CL", "St John's" : "JN"}
print(college_abbr)
college_abbr_rev = {abbr: name for name, abbr in college_abbr.items()}
print(college_abbr_rev)
college_abbr_list = sorted(college_abbr.values())
print(college_abbr_list)
```

Output:
```
{'Trinity Hall': 'TH', 'Gonville and Caius': 'CAI', 'Corpus Christi': 'CC', 'Clare': 'CL', "St John's": 'JN'}
{'TH': 'Trinity Hall', 'CAI': 'Gonville and Caius', 'CC': 'Corpus Christi', 'CL': 'Clare', 'JN': "St John's"}
['CAI', 'CC', 'CL', 'JN', 'TH']
```

#### Optional extension

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
...
```