<a href="https://colab.research.google.com/github/gsgill7/gsgill7/blob/main/Exercises/09%20Exercises.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Exercise 09.1 (checking data validity)

The Fibonacci series is valid only for $n \ge 0$. Add to the Fibonacci function in this notebook a check that raises an exception if $n < 0$. Try some invalid data cases to check that an exception is raised.

*Optional:* Use `pytest` to test that an exception *is* raised for some $n < 0$ cases.

<!-- grade_id: cell-982c3611466532a2, points: 0, grade: false, solution: true -->
```python
def f(n):
    "Compute the nth Fibonacci number using recursion"
    if n < 0:
        raise ValueError("n must be non-negative")
    if n == 0:
        return 0
    elif n == 1:
        return 1
    else:
      x1, x2 = 0, 1
      for i in range(n-1):
        x1, x2 = x2, x1 + x2
      return x2
    ...
```

<!-- grade_id: cell-9b97ea3e85d77f4a, points: 0, grade: true, solution: false -->
```python
## tests ##
#print(f(2))
# Perform some tests
assert f(0) == 0
assert f(1) == 1
assert f(2) == 1
assert f(3) == 2
assert f(10) == 55
assert f(15) == 610

# Check that ValueError is raised for n < 0
import pytest
with pytest.raises(ValueError):
    f(-1)
with pytest.raises(ValueError):
    f(-2)
```

Output:
```
1
```

## Exercise 09.2 (raising exceptions)

Modify your program from the bisection exercise in Activity 04 to raise an error if the maximum number of iterations is exceeded. Reduce the maximum allowed iterations to test that an exception is raised.

Add any other checks on the input data that you think are appropriate.

<!-- grade_id: cell-ca2c3ee62ad3b956, points: 0, grade: false, solution: true -->
```python
def my_f(x):
    """Evaluate polynomial function"""
    return x**5 / 10 + x**3 - 10 * x**2 + 4 * x + 7


def compute_root(f, x0, x1, tol, max_it):
    """Compute roots of a function using bisection"""
    error = tol + 1.0

# Iterate until tolerance is met
    it = 0
    while error > tol:
        # Compute midpoint
        x_mid = (x0 + x1) / 2

        # Evaluate function at (i) left end-point and at (ii) midpoint
        f0 = f(x0)
        f_mid = f(x_mid)

        if f_mid*f0 < 0:
          x1 = x_mid
        else:
          x0 = x_mid

        error = abs(f_mid)
        it += 1

        # Guard against an infinite loop
        if it > max_it:
          raise RuntimeError("Maximum number of iterations exceeded")

        #print(it, x_mid, error)
    return x_mid, f_mid, it
```

<!-- grade_id: cell-f6cf91ed5dd1736b, points: 0, grade: true, solution: false -->
```python
## tests ##

# Test with max_it = 30
x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=30)

# Test with max_it = 20
with pytest.raises(RuntimeError):
    x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=20)
```

Output:
```
1 1.0 2.0999999999999996
2 1.5 5.365625000000001
3 1.25 1.36669921875
4 1.125 0.4477813720703132
5 1.1875 0.4408627510070797
6 1.15625 0.008327355980872753
7 1.171875 0.21507753478363156
8 1.1640625 0.1030741602036862
9 1.16015625 0.04729774910292761
10 1.158203125 0.019466230897839054
11 1.1572265625 0.005564689502373099
12 1.15673828125 0.0013825210451265946
13 1.156982421875 0.0020907873792292975
14 1.1568603515625 0.00035405894194262544
15 1.15679931640625 0.0005142496094645566
16 1.156829833984375 8.009997302949046e-05
17 1.1568450927734375 0.00013697832466341708
18 1.1568374633789062 2.8438885864900953e-05
19 1.1568336486816406 2.5830616070976475e-05
20 1.1568355560302734 1.304116775457942e-06
21 1.156834602355957 1.2263254177469207e-05
22 1.1568350791931152 5.479569834321296e-06
23 1.1568353176116943 2.087726812760593e-06
24 1.1568354368209839 3.9180508970559913e-07
1 1.0 2.0999999999999996
2 1.5 5.365625000000001
3 1.25 1.36669921875
4 1.125 0.4477813720703132
5 1.1875 0.440862
```